# Sudoku

A “brute force” Sudoku solver tries all possible combinations by filling in blank spaces one by one. If it encounters a case that is obviously impossible, it backtracks by undoing the last step.

Some useful functions have been provided for you:

1. `generate()` to generate a puzzle.
2. `get_unused()` to get a `list` of numbers that have not yet been used.
3. `get_unfilled()` to get a `list` of coordinates that have not yet been filled.
4. `is_correct()` to check if a fully filled puzzle is correct.

Examine the autograding tests below and try out the functions to see what kind of output they provide.

```python
from sudoku import generate, get_unused, get_unfilled, is_complete, is_correct

# Autograding tests to check the above helper functions
# used = {1, 3, 7, 9}
puzzle = generate()
unused = get_unused(puzzle)
assert set(unused) == {2, 4, 5, 6, 8}, f'{set(unused)} != {{2, 4, 5, 6, 8}}'
unfilled = get_unfilled(puzzle)
assert set(unfilled) == {(0, 0), (0, 2), (1, 1), (2, 0), (2, 2)}, \
    f'{set(unfilled)} != {{(0, 0), (0, 2), (1, 1), (2, 0), (2, 2)}}'
```

## Task: Implement a recursive brute-force Sudoku solver

A recursive function `solve()` for solving a 3×3 Sudoku puzzle carries out the following steps:

1. If the puzzle is complete and correct, `return puzzle`.
2. If the puzzle is complete but incorrect, `return None`.
3. If the puzzle is still incomplete, fill in the first unused value in the first unfilled slot.
4. Solve the puzzle again (recursive call).
5. If the result is `None`, no solution is possible. Undo Step 1 and try with the next value.
6. If the result is a puzzle, return it.
7. If all values have been tried, the puzzle is unsolvable (`return None`).

Implement the above algorithm in `solve()`. Steps 1 and 2 have been done for you below.

**Hint 1:** You do not need to fill in the rest of the unfilled slots beyond the first one. The subsequent recursive calls will try other combinations for you if implemented correctly.

**Hint 2:** Your function should only have two possible return values: `puzzle`, or `None`.  
Think through your decision table: when should you return `puzzle` and when should you return `None`?

<!-- grade_id: cell-d4d13a186818b903, points: 0, grade: false, solution: true -->
```python
# The functions get_unused(), get_unfilled(), is_complete(), and is_correct()
# are available for use
from copy import deepcopy


def solve(original_puzzle):
    '''
    Attempts to solve puzzle.
    Fills the first unfilled slot with the first unused value.
    If solution check passes, calls solve() recursively with the
        new puzzle.
    Else, tries first unfilled slot with the next value.
    '''
    puzzle = deepcopy(original_puzzle)

    if is_complete(puzzle):
        if is_correct(puzzle):
            return puzzle
        else:
            return None
    else:
        unfilled = get_unfilled(puzzle)
        unused = get_unused(puzzle)
        first_unused = unused[0]
        y, x = unfilled[0]

        ### BEGIN SOLUTION ###
        all_unused_tried = False
        while not all_unused_tried:
            test_value = unused.pop(0)
            if len(unused) == 0 or unused[0] == first_unused:
                all_unused_tried = True
#             print(f'coord ({y}, {x}), trying {test_value}, remaining: {unused}')
            puzzle[y][x] = test_value
            result = solve(puzzle)
            if result is None:
                puzzle[y][x] = None
                unused.append(test_value)
            else:
#                 display(puzzle)
                return result
        # If this point is reached, all possibilities are exhausted
#         print(f'coord ({y}, {x}), all values tried. Go one level up')
        return None
        ### END SOLUTION ###
```

<!-- grade_id: cell-6fc908e5c86f2c02, points: 1, grade: true, solution: false -->
```python
# Test for correct output
puzzle = generate()
solution = solve(puzzle)
mark = '✓'
try:
    assert is_correct(solution), 'solve() does not return a correct solution'
except AssertionError as e:
    print(solution)
    mark = '✗'
    print(e)

result = f'[{mark}] Correct output  '
print(result)
```

## Extensions and Improvements

1. Write a more efficient algorithm.
2. Write a `while`-loop implementation of the 3×3 solver
3. Write a 9×9 Sudoku solver.